Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`.

Please provide **your details correctly**. I do not grade this manually and if your details are incorrect, Python script will not recognize your submission and you wil get zero for that assessment. 

```python
FIRST_NAME = "Joshua" #put your name inside the quotes
LAST_NAME = "de la Bruere" #put your name inside the quotes
WIU_ID = "919584268" #this should be available on Western Online 
```

Please ignore the cells that have 'assert' written. These will be used for autograding. 

---

### Q1 (Points 3): 

- Create a tuple object with 5 elements
- assign it to a variable name "first_tuple"
- The second element of the tuple should be a string of your choice 

<!-- grade_id: cell-ccc77de20d269c82, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
first_tuple = (1,'2',3,4,5)
```

<!-- grade_id: cell-1db38f0b321dfe9d, points: 1, grade: true, solution: false -->
```python
assert len(first_tuple) == 5
```

<!-- grade_id: cell-5876727e12450fde, points: 1, grade: true, solution: false -->
```python
assert isinstance(first_tuple, tuple) == True
```

<!-- grade_id: cell-0bd6e21bcac2f35d, points: 1, grade: true, solution: false -->
```python
assert isinstance(first_tuple[1], str) == True
```

### Q2 (Points 4 ): 

- Create a dictionary object with 4 elements
- assign it to a variable name "first_dict"
- the first two keys should be strings and the last two keys should be floats  
- the first value should be a list and the last 3 values should be integers

<!-- grade_id: cell-894a0275e6f22b95, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
first_dict ={
    "One":[ "Similar in Stucture", "JSON files can have lists", "and function with a key method as well" ],
    "two": 2, 
    3.0: 3,
    4.0: 4
}
```

<!-- grade_id: cell-d8c28fdb1a1de439, points: 1, grade: true, solution: false -->
```python
assert len(first_dict) == 4
```

<!-- grade_id: cell-1439907d93f11152, points: 1, grade: true, solution: false -->
```python
assert isinstance(first_dict, dict) == True
```

<!-- grade_id: cell-a01acc4d9e96ac7b, points: 1, grade: true, solution: false -->
```python
keys = list(first_dict.keys())
assert (isinstance(keys[0], str) and 
        isinstance(keys[1], str) and 
        isinstance(keys[2], float) and 
        isinstance(keys[3], float)) == True
```

<!-- grade_id: cell-278e51ebd00afda7, points: 1, grade: true, solution: false -->
```python
values = list(first_dict.values())
assert (isinstance(values[0], list) and 
        isinstance(values[1], int) and 
        isinstance(values[2], int) and 
        isinstance(values[3], int)) == True
```

### Q3 (Points 1): 

Convert the string 'month' in uppercase and then assign it to a variable name "upper_month"

<!-- grade_id: cell-9ddb550b9a73d91b, points: 0, grade: false, solution: true -->
```python
month = 'December'
# YOUR CODE HERE

upper_month=month.upper()
```

<!-- grade_id: cell-f9f0d0b67b5dce83, points: 1, grade: true, solution: false -->
```python
assert upper_month.isupper() == True
```

### Q4 (Points 2 ): 

- Create a variable name 'month_year' with a string 'January_2020'
- Then replace '2020' with '2021' and assign the replaced string to variable name 'updated_data'

<!-- grade_id: cell-3c14cf99b1ff1c44, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
month_year = 'January_2020'
updated_data = month_year.replace('2020', '2021')
```

<!-- grade_id: cell-ed9037a57874ed8b, points: 1, grade: true, solution: false -->
```python
assert (month_year == 'January_2020') == True
```

<!-- grade_id: cell-65a534a611cbe622, points: 1, grade: true, solution: false -->
```python
assert (updated_data[8:12] == '2021') == True
```

### Q5 (Points 5): 

**Problem: Extract words starting with character 'P' from a given string.**

Use the predefined variables **fun_string** and **final_list** (see below) and follow the four steps - each step can be converted into one line Python Statement.
- Create a list that contains the words of **fun_string** as separate elements (*Hint: string method split()*) 
- Iterate over all elements of this new list created from **fun_string** (*Hint: for loop*) 
- Check if the first element of a word is equal to uppercase string 'P' (*Hint: if statement*) 
- If the element is equal to uppercase string then append it to the predefined list **final_list** (*Hint: append method*) 

<!-- grade_id: cell-70a52b9cc4e53966, points: 0, grade: false, solution: true -->
```python
final_list = [] #This creates an empty list with no elements 
fun_string = 'Phew! Python Class is too long'

# YOUR CODE HERE
list_of_words =  fun_string.split()

for  i in fun_string.split():
    if i[0:1] =='P':
        final_list.append(i)
print(final_list)
```

<!-- grade_id: cell-a192a7acf938f69b, points: 5, grade: true, solution: false -->
```python
assert (final_list[0] == 'Phew!' and final_list[1] == 'Python') == True
```

### Q6 (Points 5): 

Create a string 'city' with the value "Chic ago" 
Extract the first element "Chic" and store it in variable "broken_string"

<!-- grade_id: cell-8fc234597af29aed, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

city = "Chic ago"
broken_string = city.split()[0]
```

<!-- grade_id: cell-1a86158f70360505, points: 2, grade: true, solution: false -->
```python
assert broken_string == "Chic"
```

### Q7 (Points 5): 

- Iterate over all elements of a sequence of numbers between 0 to 100
- If the element is divisible by 5 then append it to the predefined list **final_numbers** 
(*Hint: to check the divisibility we want the remainder to be zero. We learned that '%' is the remainder operator*) 

<!-- grade_id: cell-d12ce62fc1f27814, points: 0, grade: false, solution: true -->
```python
final_numbers = [] #This creates an empty list with no elements 
# YOUR CODE HERE

for i in range(0, 100,1):
    if i % 5 == 0:
        final_numbers.append(i)
    
```

<!-- grade_id: cell-446b95a494003eac, points: 3, grade: true, solution: false -->
```python
assert len(final_numbers) != 0
for item in final_numbers:
    assert (item%5 == 0)
```

```python

```