INSTRUCTIONS:

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`.

Please provide **your details correctly**. 

Please replace the **id** in the file name with your WIU_ID.

Please do not change the name of the file. 

I do not grade this manually and if your details are incorrect, Python script will not recognize your submission and you wil' get zero for that assessment. 

```python
FIRST_NAME = "Joshua" #put your name inside the quotes
LAST_NAME = "de la Bruere" #put your name inside the quotes
WIU_ID = "919584268" #this should be available on Western Online 
```

Please ignore the cells that have 'assert' written. These will be used for autograding. 

---

## Q1 (Points 2): 

- Create a function named 'fav_drink_food'
- It can take arbitrary number of keyword arguments using **kwargs
- For this specific use case - it take two keyword arguments 
    1. 'summer_food'
    2. 'winter_food' 
    
- The function returns the string:

    'My favorite summer food is `summer_food`, and winter food is `winter_food`'

    - The `summer_food` will be the value of the keyword summer_food 

    - The `winter_food` will be the value of the keyword winter_food 


<!-- grade_id: cell-ccc77de20d269c82, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

def fav_drink_food (summer_food="none", winter_food="none", **kwargs ):
    string = 'My favorite summer food is '+ summer_food + ', and winter food is ' + winter_food
    return(string)
```

```python
#test your output here 
fav_drink_food(winter_food='Pasta', summer_food='Water Melon') 
```

Output:
```
'My favorite summer food is Water Melon, and winter food is Pasta'
```

<!-- grade_id: cell-1db38f0b321dfe9d, points: 2, grade: true, solution: false -->
```python
#Worth 2 points
assert fav_drink_food(winter_food='Pasta', summer_food='Water Melon') \
        == 'My favorite summer food is Water Melon, and winter food is Pasta'

assert fav_drink_food(winter_food='Meat Loaf', summer_food='Soup') \
        == 'My favorite summer food is Soup, and winter food is Meat Loaf'
```

# Q2 (Points 5): 

- Create a function named 'cal_average'
- It take a list of numbers 
- Return the average of the list 

*Hint:*
1. use `sum(list_name)` to calculate the sum, as we did in lectures
2. use `len(list_name)` to calculate size or length or number of observations


<!-- grade_id: cell-894a0275e6f22b95, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def cal_average(list_name):
    average=sum(list_name)/len(list_name)
    return(average)
```

```python
#test your output here 
cal_average([3,4,5])
```

Output:
```
4.0
```

<!-- grade_id: cell-d8c28fdb1a1de439, points: 5, grade: true, solution: false -->
```python
assert cal_average([2,3]) == 2.5
assert cal_average([3,4,5]) == 4.0
```

## Q3 (Points 8): 

- Create a function named 'cal_variance'
- It take a list of numbers 
- Return the variance of the list 

*Hint:* You may follow these steps
1. calculate the average
2. calculate the squared deviation for each obseration 
    - create an empty list let's say numerator
    - iterate over each item in the list of numbers 
    - subtract that item from the average to get the deviation
    - square the deviation 
    - append this deviation to the list numerator at each run of for loop
3. calculate the length of the list 
4. divide the sum of 'numerator' by lenght to get the variance

<!-- grade_id: cell-780ebe3021c1792b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def cal_variance (list_num):
    
    # Get the average
    average = cal_average(list_num)
    
    #empty list for deviation
    deviation = []
    
    # for loop to get deviation
    for i in range(0, len(list_num),1):
        deviation.append((average-list_num[i])**2)
    
    # Varience calculation
    varience = sum(deviation)/(len(deviation)-1)
    
    # variables out (average, length, sum, and varience)
    var_out = [average, len(list_num), sum(list_num), varience]
    return(var_out)

```

```python
#test your answer here 
cal_variance([3,4,5])
```

Output:
```
[4.0, 3, 12, 1.0]
```

<!-- grade_id: cell-ffad039c5775b759, points: 2, grade: true, solution: false -->
```python
#Worth 2 points
assert cal_variance([3,4,5])[0] == 4.0 #average
```

<!-- grade_id: cell-bddb923090daf0b9, points: 2, grade: true, solution: false -->
```python
#Worth 2 points
assert cal_variance([3,4,5])[1] == 3 #len
```

<!-- grade_id: cell-073a22d9ca6d7120, points: 2, grade: true, solution: false -->
```python
#Worth 2 points
assert cal_variance([3,4,5])[1] == 3 #sum(numerator)
```

<!-- grade_id: cell-c50f338883768425, points: 2, grade: true, solution: false -->
```python
#Worth 2 points
assert cal_variance([3,4,5])[3] == 1.0 #variance
```

# Q4 (Points 2): 

- Create a function named 'print_boolean'
- It takes an integer or a float as an argument
- The function returns the string:
    - 'It should be sunny outside' if the value passed is greater than or equal to 32
    - 'It is cold' if the value passed is less than 32

<!-- grade_id: cell-c211e86770d8d619, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def print_boolean (temp):
    if(temp >= 32):
        return('It should be sunny outside')
    return('It is cold')
```

```python
# test your output here 
print_boolean(34)
```

Output:
```
'It should be sunny outside'
```

<!-- grade_id: cell-7ce9207fc89f964f, points: 1, grade: true, solution: false -->
```python
#Worth 1 points
assert print_boolean(45) == 'It should be sunny outside'
```

<!-- grade_id: cell-bf04d670e3b2184a, points: 1, grade: true, solution: false -->
```python
#Worth 1 points
assert print_boolean(30) == 'It is cold'
```

## Q2 (Points 3): 

- Create a function named 'cal_status'
- It can take arbitrary number of positional arguments using *args
- The function returns the string:
    - 'All numbers are negative' if all values passed are negative
    - 'Positive and negative numbers are there' is there are positive and negative values 
    - 'All numbers are positive' if all values passed are positive

<!-- grade_id: cell-3f87113ff48074fb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def cal_status (*numbers):
    positive = False
    negative = False
    for i in numbers:
        if i > 0:
            positive = True
        elif i< 0:
            negative = True
        else:
            return('Positive and negative numbers are there' )
    if positive == True:
        if negative == True:
            return('Positive and negative numbers are there' )
        else:
            return('All numbers are positive')
    else:
        return('All numbers are negative')
        
```

```python
# test your output here 
cal_status(1,2)
```

Output:
```
False
True
'All numbers are positive'
```

<!-- grade_id: cell-4ac1e7f4c02b9d81, points: 1, grade: true, solution: false -->
```python
#Worth 1 points
assert cal_status(-2, 3, 4) == 'Positive and negative numbers are there'
```

<!-- grade_id: cell-8045c781c7a3ca5f, points: 1, grade: true, solution: false -->
```python
#Worth 1 points
assert cal_status(-2, -3, -4, -9) == 'All numbers are negative'
```

<!-- grade_id: cell-1ddf6e0721bf4bf1, points: 1, grade: true, solution: false -->
```python
#Worth 1 points
assert cal_status(2, 3, 4, 3, 9) == 'All numbers are positive'
```